Time to Eat - A Personalized Circadian Eating Schedule Leads to Weight Loss Without Imposing Calorie Restriction: A Randomized Controlled Pilot Study

Author response:

The following is the authors’ response to the previous reviews.

We would like to remind the editors and reviewers that the present project is a pilot study that does not claim to produce definitive results. Pilot studies are exploratory preliminary studies to test the validity of hypotheses, the feasibility of a study as well as the research methods and the study design. From our point of view, our hypotheses and the feasibility of the pilot study have been confirmed to such an extent that the implementation of a larger study is justified. At the same time, it became clear during the pilot that the methods and design need to be adapted in some areas in order to increase the reliability of the results - a finding that pilot studies are usually conducted to obtain. We discussed these limitations in detail in order to explain the planned changes in the follow-up study. What the reviewers and editors interpret as incompleteness is therefore due to the nature of a pilot study. We consider it necessary that appropriate standards are taken into account in the evaluation of the present work.

In addition, we would like to make a counterstatement as to what our main claims, which should be used to assess the strength of evidence, are - and what they are not:

In the introduction, we describe the background that led to the formation of our hypotheses: Previous animal and human studies show that food, along with light, serves as the main Zeitgeber for circadian clocks. It has also been shown that chrononutrition can lead to weight loss and improved well-being. Based on this, we hypothesized that individualized meal timing can enhance these positive effects. This hypothesis has been validated on the basis of the available results. Contrary to what the editors and reviewers stated, the assumption that the observed beneficial effects are indeed related to an alteration or resetting of endogenous circadian rhythms was not intended to be investigated in this study and is not one of our main claims. This has already been sufficiently demonstrated and, in our view, need not and should not be repeated in every study on chrononutrition. Accordingly, this assumption was not formulated as a working hypothesis or main claim. It is described in the paper as a potential mechanism, the assumption of which is justified on the basis of previous studies. The lack of a corresponding examination and the erroneous insinuation that corresponding results were nevertheless listed by us in the paper as a main claim should therefore not be used as a criterion for downgrading the assessment of the strength of evidence.

The main criticism of our study is the collection of data using self-reported food and food quantities. This form of data collection is indeed prone to error, as there is little control over the accuracy of the reported data. However, we believe that this problem is limited in scope.

(1) Contrary to what the editors and reviewers claim, at no point do we write that we are convinced that food intake has not changed. On the contrary, in Figure 2 we explicitly show that there was a change in what some participants reported to us regarding their food intake. We make it clear throughout the text that we could not find any correlation between weight change and the changes in the reports of food quantities/meals. These statements are correct and only what are actual and formulated main claims should be included in the evaluation of the study.

(2) As previously stated, we conducted analyses that suggest that an unreported reduction in food intake is unlikely to be the cause of weight loss. For the most part, participants did not change their reporting behavior during the exploration and intervention phases. That is, participants who underreported food intake reported similar amounts in both phases of the study, but lost weight only in the intervention phase. To explain their weight loss with imprecise reporting, it would have to be assumed that these participants began to eat less in the intervention phase and at the same time report more in order to achieve similar calorie counts and food composition in the evaluation. We consider such behavior to be very unlikely, especially since it would apply to numerous participants.

(3) The editors and reviewers reduce the results to the absence of a correlation between weight loss and reported food quantity and composition. In their assessment of the significance of the findings, however, they ignore the fact that we did find a significant correlation in our analyses, namely between weight loss and an increase in the regularity of food intake. There is no correlation between an increase in regularity and a reduction in reported calories (R2 = 0.01472). This is credible in our view, as it is unlikely that the more regularly participants ate, the more pronounced the error in their reports was (while in reality they ate less than before).

(4) We also had the requirement for the study design that the participants could carry out the intervention in their normal everyday life and environment in order to test and ensure implementation in real life. We consider it unrealistic to be able to monitor food intake continuously and without interruption over a period of several weeks under these conditions. We therefore see no alternative to self-reporting. As the reviewers and editors did not suggest any alternative methods of data collection that would fulfil the requirements of our study, we assume that, despite criticism and reservations, they generally agree with our assessment and take this into account in their evaluation.

It is still criticized that some confounding factors are present. The reviewer makes no reference to the fact that we either eliminated these in the last version submitted (age range), identified them as unproblematic (unmatched cohorts, menstrual cycle, shift work) or even deliberately used them in order to be able to test our hypothesis more validly (inclusion of individuals with normal weight, overweight, and obesity).

Besides, the use of actimeters to determine circadian rhythms as proposed by the editors and reviewers is not valid for this study and the requirement to use them to determine a circadian reset in the eLife assessment is misleading and inappropriate. This instrument only measures physical activity, but not the physiological parameters that are relevant for an investigation in this field of research.

For the assessment of chronotype alone, the MCTQ questionnaire is a valid instrument that has been validated several times against actimetry (e.g., DOIs: 10.1080/07420528.2022.2025821, 10.1080/07420528.2023.2202246, 10.1016/j.ijpsycho.2016.07.433, 10.1155/2018/5646848). The reviewer's statement that the MCTQ questionnaire is unreliable for determining chronotype is unsupported and incorrect.

Equally unproven is the statement that any form of imposed diet appears to lead to weight loss over a period of several months.

Nevertheless, in order to prevent further misunderstandings, we have revised our text in a number of places and clarified that our statements are not irrefutable assertions, but potential interpretations of the results obtained in the pilot study, which are to be analyzed in more detail with regard to the planned more comprehensive study.

Changes from the earlier version to the revised version

Abstract

Abstract

Objective

Earlier version: Many weight loss strategies are based on the restriction of calories or certain foods. Here, we tested a weight loss intervention based solely on increasing the regularity of meals, presuming that this allows the circadian system to optimally prepare the food metabolism for these times.

Revised version: Many weight loss strategies are based on the restriction of calories or certain foods. In this pilot study, we preliminary tested a weight loss intervention based solely on increasing the regularity of meals. The assumption is based on the evidence that eating at fixed times allows the circadian system to optimally prepare the food metabolism for these times.

Participants & Methods

Earlier version: In a two-group, single center randomized-controlled single-blind pilot study (pre-registration DRKS00021419) with participants aged 18-65 years and BMI ≥ 22 kg/m², we used a smartphone application to identify the times at which each participant eats particularly frequently and asked participants of the experimental group to restrict their meals to only these times for six weeks. Control participants received sham treatment. Primary outcome was self-reported body weight/BMI and secondary outcome the well-being of participants.

Revised version: In a two-group, single center randomized-controlled single-blind pilot study (pre-registration DRKS00021419) with participants aged 18-65 years and BMI ≥ 22 kg/m2, we used a smartphone application to identify the times at which each participant eats particularly frequently and asked participants of the experimental group to restrict their meals to only these times for six weeks. Control participants received sham treatment. Primary outcome was body weight/BMI and secondary outcome the well-being of participants.

Results

Earlier version: Of 148 participants entering the pilot study, 121 were randomized and of these 100 (control: 33, experimental: 67) finished the study. Our results show that the more regular the meals of participants of the experimental group became, the more weight/BMI they lost, averaging 2.62 kg (0.87 kg/m²); p < 0.0001 (BMI: p < 0.0001) compared to an insignificant weight loss of 0.56 kg (0.20 kg/m²) in the control group; p = 0.0918 (BMI: p = 0.0658). Strikingly, weight loss was not related to changes in self-reported calories, food composition, and other food-related factors. Additionally, physical and mental well-being improved significantly.

Revised version: Of 148 participants taking part in the pilot study, 121 were randomized, of whom 100 (control: 33, experimental: 67) completed the study. Our results show that the more regular the meals of participants of the experimental group became, the more weight/BMI they lost, averaging 2.62 kg (0.87 kg/m2); p < 0.0001 (BMI: p < 0.0001) compared to an insignificant weight loss of 0.56 kg (0.20 kg/m2) in the control group; p = 0.0918 (BMI: p = 0.0658). Strikingly, weight loss was not related to self-reported changes in calories, food composition, and other food-related factors. Additionally, physical and mental well-being improved significantly.

Conclusion

Earlier version: In summary, increasing the regularity of meals causes participants to lose excess body weight and improves overall well-being.

Revised version: In summary, increasing the regularity of meals may cause participants to lose excess body weight and improve overall well-being. These promising results justify a larger-scale study, albeit with a more rigorous study design.

Highlights

Individual optimal times for meals are determined via an app-based meal diary.

Generation of a structure plan for mealtimes are adjusted to individual circadian clocks.

Following this plan, participants lost an average body weight of 2.6 kg over six weeks.

Earlier version: Weight loss is achieved without changes in self-reported food quantity or composition.

Revised version: Weight loss does not correlate with self-reported changes in food quantity or composition.

Earlier version: Regular mealtimes contribute to the improvement of the general well-being.

Revised version: Regular mealtimes contribute to the improvement of general well-being.

Introduction

Circadian clocks induce endogenous 24-hour rhythms in the expression of more than half of all genes across all tissues [1, 2], allowing the body to anticipate daily changes between day and night and to prepare and harmonize physiological processes accordingly. However, this requires exposure to very regular environmental time cues, so-called Zeitgebers, such as food. Thus, on the one hand, the circadian system contributes to the optimization of food metabolism; on the other hand, the intake of food serves as a potent Zeitgeber. Both mouse and human studies have shown the negative consequences of circadian disturbances on metabolism, body weight [3–11], and also mental health [12–14], and demonstrated that restriction of meals to certain time spans of the day can counteract these harmful effects [15–19]. In the recent past, some human studies have shown that restricting food intake to certain times of the day (time-restricted eating, TRE / time-restricted feeding, TRF) can indeed lead to weight loss. In most cases, the interventions were successful, and participants lost ∼0.9 - 3.5 kg body weight within 12 weeks when they reduced their food intake to 8–10-hour intervals per day [19–22]. Furthermore, they suggest that TRE is more successful when the meal window is earlier in the day rather than later in the day [23]. However, previous human studies merely established rough periods during which participants could still eat irregularly.

Earlier version: It has previously been shown that the success of a hypocaloric diet can be increased if the distribution of daily calories consumed is adapted to the individual chronotypes in such a way that morning types consume many calories in the first half of the day and evening types consume many calories in the second half of the day [24]. We expect that the success of previous chrononutrition interventions can be further improved by determining each participant’s eating time profile and creating a personalized meal plan accordingly, even if participants continue to consume their usual amounts and foods. We are testing this in the present pilot study, which is intended to examine in principle whether this assumption and possibly a more in-depth continuation of the study is justified.

Revised version: It has previously been shown that the success of a hypocaloric diet can be increased if the distribution of daily calories consumed is adapted to the individual chronotypes in such a way that morning types consume many calories in the first half of the day and evening types consume many calories in the second half of the day [24]. We expect that the success of previous chrononutrition interventions can be further improved by determining each participant’s eating time profile and creating a personalized meal plan accordingly, even if participants are instructed to continue consuming their usual amounts and types of foods. We are testing this in the present pilot study, which is intended to examine in principle whether this assumption and possibly a more in-depth continuation of the study is justified.

Earlier version: Therefore, the primary aim is to individually determine and set optimal times for food intake and to correlate the reduction of mealtime variability with alterations in body weight. We include participants of Body Mass Index (BMI) classifications normal to extremely obese to test whether participants who exceed normal BMIs particularly benefit from the intervention. Because of the impact of circadian clocks on other body functions, we further hypothesize that regular food intake contributes to an overall improvement in well-being, such as the subjective feeling of general health, sleep quality, affective state, and self-efficacy. After screening, recruitment, a session to introduce the content and technical aspects of the study, and the completion of questionnaires on well-being, the eating times and food composition of all participants were recorded for two weeks using a smartphone application (Fig. S1A). Based on this data, a meal schedule was designed for each participant individually, which the experimental participants were asked to follow for the next six weeks while data on eating times and food composition continued to be collected. For control subjects, we specified a long 18-h window during which they could continue to eat irregularly. The same questionnaires as at the beginning were then collected again. The participants then had the opportunity to end the intervention or to continue without further data collection. In both cases, however, an attempt was made to reach them for a follow-up assessment. Since the program focuses on eating times and, unlike many other dietary programs, explicitly does not aim to limit calories or certain types of food, we have named it "Time to Eat" (Fig. S1B).

Revised version: Therefore, the primary aim is to individually determine and set optimal times for food intake and to correlate the reduction of mealtime variability with alterations in body weight. We include participants of Body Mass Index (BMI) classifications normal to extremely obese to test whether participants who exceed normal BMIs particularly benefit from the intervention. Because of the impact of circadian clocks on other body functions, we further hypothesize that regular food intake contributes to an overall improvement in well-being, such as the subjective feeling of general health, sleep quality, affective state, and self-efficacy. After screening, recruitment, a session to introduce the content and technical aspects of the study, and the completion of questionnaires on well-being, all participants recorded their eating times and food composition for two weeks using a smartphone application (Fig. S1A). Based on this data, a meal schedule was designed for each participant individually, which the experimental participants were asked to follow for the next six weeks while data on eating times and food composition continued to be collected. For control subjects, we specified an 18-h window during which they could continue to eat irregularly. The same questionnaires as at the beginning were then collected again. The participants then had the opportunity to end the intervention or to continue without further data collection. In both cases, however, an attempt was made to reach them for a follow-up assessment after four weeks. Since the program focuses on eating times and, unlike many other dietary programs explicitly does not aim to limit calories or certain types of food, we have named it “Time to Eat” (Fig. S1B).

Results

Participants

For this pilot study, a total of 148 participants were recruited. Of these, 121 were randomly allocated to the experimental group (EG) (n=79) and the control group (CG) (n=42) (Fig. S1C), which had non-significantly different mean BMIs of 25.9 (SD: ± 4.138) (CG) and 27.5 (SD: ± 4.627) (EG) (Tab. S1). At the end of the intervention (T2), data were available from 100 participants (EG: n=67, CG: n=33). Their baseline data indicate that participants in both groups were, on average, living with overweight according to BMI classification, with the included CG participants having a BMI 25.5 (SD: ± 3.827) and being slightly less overweight, on average, than the analyzed EG participants with a BMI of 27.4 (SD: ± 4.616) (Tab. S1). There were no significant BMI differences between completers and drop-outs in either group (Fig. S2A, Tab. S2). Reasons for drop-outs were, according to self-report, illness/accident during the study and unwillingness or failure to adhere to fixed eating times.

Improvement of eating time regularity

All study participants were asked to record each occasion when anything containing calories, such as meals, snacks, sugared or milk-containing drinks, was consumed during a 14-day exploration and a six-week intervention phase (Fig. 1A, B, S1C). During the exploration and intervention phases, a total of 13,838 and 34,564 calorie intake occasions were recorded from completers, respectively. Most EG participants were assessed as having three eating times (n=53), some reported eating four meals (n=13), and one had two main meals per day. The average MTVS for breakfast, lunch, dinner, and total meals of the exploration phase of both groups was around 4, which corresponds to a daily deviation of ∼120 min for each meal. In the EG group, mean scores improved to less than to 2 during the intervention phase, corresponding to a mean deviation of less than ± 30min (Fig. 1C, D, Tab. S3). In contrast, in the CG group, there was no improvement but some significant increase of the MTVS.

Added in the revised version: Improvement in regularity of meals, body weight changes, and correlation between loss of body weight and regularity of eating meals.

Earlier version: Improvement in regularity of meals, body weight changes, and correlation between loss of body weight and regularity of eating meals. (A, B) Times of day of calorie intake occasions of 15 representative participants of the EG (A) and the CG (B) during the two-week exploration phase (left) and the six-week intervention phase (right). Each point represents a calorie intake occasion. (C, D) Change in MTVS of participants in the CG (gray) and EG (cyan) of breakfast, lunch, and dinner (C) and of all meals combined (D), including snacks, if snacks were taken. 2-way repeated-measures ANOVA with Bonferroni post hoc test comparing CG and EG during the exploration or intervention phase and T0-T1 and T1-T2 within groups, * p ≤ 0.05, ** p ≤ 0.01, **** p ≤ 0.0001. CG: n=33; EG: n=64-67, missing n-values in the EG are due to the different number of meals for each participant. Detailed statistics in Tables S3. (E) Within group differences of body weight during the study. Body weight was collected at different time points, at the beginning of the study (T0), after the exploration phase (T1), after the intervention phase (T2), and after the follow-up phase (T3). During the follow-up phase, 34 EG participants voluntarily continued the intervention (continuing; black line and open circle) while 33 participants no longer followed the intervention (discontinuing; gray, dashed line and open circle). Only three control subjects continued their intervention in the follow-up phase, which is why a division into continuing and discontinuing was omitted in this group. Data are normalized so that all baseline data at T0 equals 0 to represent change. Differences of all CG and EG participants between T0, T1, and T2 were calculated with a 1-way ANOVA, as were differences between T2 and T3 of CG participants (Bonferroni post-hoc test, **** p ≤ 0.0001). Differences between T2 and T3 of continuing and discontinuing EG participants were calculated with a paired t-test °° p ≤ 0.01, °°°° p ≤ 0.0001. Detailed statistics in Table S4. (F) Between group differences of BMI from T0-T3. Mixed-effect model with Bonferroni post-hoc test comparing CG and EG at different phases during the study, **** p ≤ 0.0001. Detailed statistics in Table S5. (G) Significant correlation between body weight/BMI and MTVS. The more participants improved their regularity of meals (decreasing MTVS), the more weight/BMI they lost. Linear regression, CG: n=33; EG: n=67, ** p ≤ 0.01. Detailed statistics in Tab. S6. (H) Significant correlation between body weight/BMI and MTVS of lunch and dinner. The more participants improved their regularity of lunch and dinner (decreasing MTVS), the more weight/BMI they lost. Linear regression, CG: n=33; EG: n=67, ** p ≤ 0.01. Detailed statistics in Tab. S6. (I) Significant correlation between change in BMI, baseline BMI, and improvement in MTVS. Individuals with higher BMI benefit more from similar improvements in MTVS (shades of data points) than individuals with lower BMI in EG participants. Multiple regression, n=67, * p ≤ 0.05, ** p ≤ 0.01. Detailed statistics in Tab. S8.

Revised version: (A, B) Times of day of calorie intake occasions of 15 representative participants of the EG (A) and the CG (B) during the two-week exploration phase (left) and the six-week intervention phase (right). Each point represents a calorie intake occasion. (C, D) Change in MTVS of participants in the CG (gray) and EG (cyan) of breakfast, lunch, and dinner (C) and of all meals combined (D), including snacks, if snacks were taken. 2-way repeated-measures ANOVA with Bonferroni post hoc test comparing CG and EG during the exploration or intervention phase and T0-T1 and T1-T2 within groups, * p ≤ 0.05, ** p ≤ 0.01, **** p ≤ 0.0001. CG: n=33; EG: n=64-67, missing n-values in the EG are due to the different number of meals for each participant. Detailed statistics in Tables S3. (E) Within group differences of body weight during the study. Body weight was collected at different time points, at the beginning of the study (T0), after the exploration phase (T1), after the intervention phase (T2), and after the follow-up phase (T3). During the follow-up phase, 34 EG participants voluntarily continued the intervention (continuing; black line and open circle) while 33 participants no longer followed the intervention (discontinuing; gray, dashed line and open circle). Only three control subjects continued their intervention in the follow-up phase, which is why a division into continuing and discontinuing was omitted in this group. Data are normalized so that all baseline data at T0 equals 0 to represent change. Differences of all CG and EG participants between T0, T1, and T2 were calculated with a 1-way ANOVA, as were differences between T2 and T3 of CG participants (Bonferroni post-hoc test, **** p ≤ 0.0001). Differences between T2 and T3 of continuing and discontinuing EG participants were calculated with a paired t-test °° p ≤ 0.01, °°°° p ≤ 0.0001. Detailed statistics in Table S4. (F) Between group differences of BMI from T0-T3. Mixed-effect model with Bonferroni post-hoc test comparing CG and EG at different phases during the study, **** p ≤ 0.0001. Detailed statistics in Table S5. (G) Significant correlation between body weight/BMI and MTVS. The more participants improved their regularity of meals (decreasing MTVS), the more weight/BMI they lost. Linear regression, CG: n=33; EG: n=67, ** p ≤ 0.01. Detailed statistics in Tab. S6. (H) Significant correlation between body weight/BMI and MTVS of lunch and dinner. The more participants improved their regularity of lunch and dinner (decreasing MTVS), the more weight/BMI they lost. Linear regression, CG: n=33; EG: n=67, ** p ≤ 0.01. Detailed statistics in Tab. S6. (I) Significant correlation between change in BMI, baseline BMI, and improvement in MTVS. Individuals with higher BMI benefit more from similar improvements in MTVS (shades of data points) than individuals with lower BMI in EG participants. Multiple regression, n=67, * p ≤ 0.05, ** p ≤ 0.01. Detailed statistics in Tab. S8.

Weight change

Earlier version: During the exploration phase, when no instructions on eating times were given, CG and EG participants had a statistically non-significant average loss of 0.17 kg (0.06 kg/m²) and 0.21 kg (0.06 kg/m²), respectively (Figs. 1E, S2B-F, Tabs. S4, S5). However, during the intervention phase, when EG participants took their meals more regularly, they lost an average of about 2.6 kg (95% CI [-2.906, -1.915]) and 0.81 kg/m² (95% CI [-0.9729, -0.6420]) (Figs. 1E, S2B-E, Tabs. S4, S5), which translates to an average weight loss of 0.40 kg (0.145 kg/m²) per week during the intervention. The weight loss observed in this group is independent of the age of the participants and their sex (Fig. S2G, Tabs. S6, S2). In contrast, weight and BMI changes of -0.39 kg (95% CI [-0.9125, 0.1307]) and -0.14 kg/m² (95% CI [-0-3152, 0.0304]), respectively, were not significant in the CG (Figs. 1E, S2B, C, F, Tabs. S4, S5), resulting in a significant difference in weight loss between groups at the end of the intervention (Figs. 1F, S2C, Tab. S5). The inclusion of the last available weight and BMI data of the drop-outs in an end-point analysis did not lead to any notable change in this result (Fig. S2H, Tab. S5).

Revised version: During the exploration phase, when no instructions on eating times were given, CG and EG participants had a statistically non-significant average loss of 0.17 kg (0.06 kg/m2) and 0.21 kg (0.06 kg/m2), respectively (Figs. 1E, S2B-F, Tabs. S4, S5). However, during the intervention phase, when EG participants took their meals more regularly, they lost an average of about 2.6 kg (95% CI [−2.906, −1.915]) and 0.81 kg/m2 (95% CI [−0.9729, −0.6420]) (Figs. 1E, S2B-E, Tabs. S4, S5), which translates to an average weight loss of 0.40 kg (0.145 kg/m2) per week during the intervention. The weight loss observed in this group is independent of the age of the participants and their sex (Fig. S2G, Tabs. S6, S2). In contrast, weight and BMI changes of −0.39 kg (95% CI [−0.9125, 0.1307]) and −0.14 kg/m2 (95% CI [−0-3152, 0.0304]), respectively, were not significant in the CG (Figs. 1E, S2B, C, F, Tabs. S4, S5), resulting in a significant difference in weight loss between groups at the end of the intervention (Figs. 1F, S2C, Tab. S5). The inclusion of the last available weight and BMI data of the drop-outs in an end-point analysis did not lead to any notable change in this result (Fig. S2H, Tab. S5).

Earlier version: Four weeks after the intervention ended, all participants were contacted again for a follow-up assessment. 34 of the 67 EG participants reported to have maintained the regularity of their meals and continued to lose body weight significantly, on average another 1.12 kg (95% CI [-1.479, -0.7626]) and 0.36 kg/m² (95% CI [-0.4907, -0.2457]), respectively (Figs. 1E, S2B, D, E, Tab. S4). In contrast, participants who reported to have resumed irregular eating gained a significant average of 0.48 kg (95% CI [0.1984, 0.7616]) and 0.17 kg/m² (95% CI [0.07434, 0.2603]), respectively (Fig. 1E, S2B, D, E, Tab. S4).

Revised version: Four weeks after the intervention ended, all participants were contacted again for a follow-up assessment. 34 of the 67 EG participants reported to have maintained the regularity of their meals and continued to lose body weight significantly, on average another 1.12 kg (95% CI [−1.479, −0.7626]) and 0.36 kg/m2 (95% CI [−0.4907, −0.2457]), respectively (Figs. 1E, S2B, D, E, Tab. S4). In contrast, participants who reported to have resumed irregular eating gained a significant average of 0.48 kg (95% CI [0.1984, 0.7616]) and 0.17 kg/m2 (95% CI [0.07434, 0.2603]), respectively (Fig. 1E, S2B, D, E, Tab. S4).

Relation between weight change and improvement in the regularity of meals

To examine the extent to which EG participants’ weight loss was effectively related to improvements in meal regularity, we correlated changes in body weight and BMI with individual MTVS changes. Consistent with our hypothesis, a significant relationship exists between weight loss and improvement in regularity in the EG (Fig. 1G, Tab. S6). A breakdown of meals shows that increasing the regularity of lunch and dinner has the greatest effects, with no significant relationship between increasing the regularity of breakfast and weight loss (Fig. 1H, Tab. S6). In contrast, in the CG, there were no or little improvements in regularity and, accordingly, no correlation with weight/BMI changes.

Further analyses revealed that all MTVS up to 3, which corresponds to a time window for meals of up to ±45 min, resulted in comparably effective weight loss, whereas the efficiency of weight loss decreases significantly with an MTVS greater than 3 (Tab. 1).

Relation between weight change and baseline BMI

Earlier version: A weight loss intervention which is particularly based on metabolic optimization is likely to be most effective for individuals with exceeded body weight. In fact, the higher the baseline BMI of EG participants at T1, the greater the weight loss (Fig. S2I, Tab. S6). Stratifying the subjects by baseline BMI, including the control group, additionally shows that in all strata the BMI of EG subjects decreased more than that of CG subjects, even if they had the same initial BMI (Fig. S2J, Tab. S7). Multiple regression shows that the reduction in MTVS and baseline BMI together predict the weight loss of the participants (Fig. 1I, Tab. S8). These data show that regularity in meal intake can lead to a significant loss of body weight and suggest that this effect is especially pronounced in individuals with a higher BMI.

Revised version: A weight loss intervention which is particularly based on metabolic optimization is likely to be most effective for individuals with exceeded body weight. In fact, the higher the baseline BMI of EG participants at T1, the greater the weight loss (Fig. S2I, Tab. S6). Stratifying the subjects by baseline BMI, including the control group, additionally shows that in all strata the BMI of EG subjects decreased more than that of CG subjects, even if they had the same initial BMI (Fig. S2J, Tab. S7). Multiple regression shows that the reduction in MTVS and baseline BMI together predict the weight loss of the participants (Fig. 1I, Tab. S8). These data indicate that regularity in meal intake can lead to a significant loss of body weight and suggest that this effect is especially pronounced in individuals with a higher BMI.

Relation between weight change and self-reported food quantity and composition

Earlier version: By restricting meals to specific times and shortening the time window in which food is eaten, it is feasible that EG participants ate less or changed food composition during the intervention phase and therefore lost weight. However, on average, the reported food energy ingested by CG and EG participants did not substantially change during the study (Fig. 2A, S2K, L, Tab. S3, 6), apart from a minor average reduction of ∼350 kJ (95% CI [-627.6, -67.67]) per day in the EG (Fig. S2K, Tab. S3), and did not differ between CG and the EG. As they are rather small, these differences do not add up to a significant reduction in the total average energy consumed over the course of the study. Accordingly, the determined cumulative average energy intake during the intervention phase did not differ on any day from the estimated energy intake based on an extrapolation of the cumulative food consumption during the exploration phase, neither in the CG nor in the EC group (Fig. 2B). Nevertheless, there were some participants who reported having consumed fewer calories during the intervention phase, while some also reported having ingested more calories. However, consistent with our hypothesis that high regularity of mealtimes optimizes metabolism of similar amounts of food, there is no relationship between reported changes in average daily and cumulative energy intake and change in body weight/BMI of either CG or EG participants (Fig. 2C, D, Tab. S6). Likewise, the composition of the ingested foods did not change considerably in either group, and participants reported a normal distribution of the macronutrients of fat, carbohydrate, and protein in both phases of the study (Fig. 2E, S2M, Tab. S3). Also, there was no correlation between weight change and change in reported average daily and cumulative macronutrient intake (Fig. 2F, S2N, Tab. S6).

Revised version: By restricting meals to specific times and shortening the time window in which food is eaten, it is feasible that EG participants ate less or changed food composition during the intervention phase and therefore lost weight. However, on average, the reported food energy ingested by CG and EG participants did not substantially change during the study (Fig. 2A, S2K, L, Tab. S3, 6), apart from a minor average reduction of ∼350 kJ (95% CI [−627.6, −67.67]) per day in the EG (Fig. S2K, Tab. S3), and did not differ between CG and the EG. As they are rather small, these differences do not add up to a significant reduction in the total average energy consumed over the course of the study. Accordingly, the determined cumulative average energy intake during the intervention phase did not differ on any day from the estimated energy intake based on an extrapolation of the cumulative food consumption during the exploration phase, neither in the CG nor in the EC group (Fig. 2B). Nevertheless, there were some participants who reported having consumed fewer calories during the intervention phase, while some also reported having ingested more calories. However, consistent with our hypothesis that high regularity of mealtimes optimizes metabolism of similar amounts of food, there is no relationship between reported changes in average daily and cumulative energy intake and change in body weight/BMI of either CG or EG participants (Fig. 2C, D, Tab. S6). Likewise, the composition of the reported ingested foods did not change considerably in either group, and participants reported a normal distribution of the macronutrients of fat, carbohydrate, and protein in both phases of the study (Fig. 2E, S2M, Tab. S3). Also, there was no correlation between weight change and change in reported average daily and cumulative macronutrient intake (Fig. 2F, S2N, Tab. S6).

Added in the revised version: Changes in body weight/BMI are not related to changes in reported caloric intake and composition of macronutrients.

Earlier version: Changes in body weight/BMI are not related to changes in reported caloric intake and composition of macronutrients. Changes in reported daily caloric intake (A), cumulative caloric intake (B) and food composition (E) and their relationship to weight/BMI changes (C, D, and F). Data in (A) are normalized so that all baseline data of the exploration phase equal 1. To represent the x-fold change, data from the intervention phase are divided by data from the exploration phase of the corresponding participant. (A, E) 2-way repeated-measures ANOVA with Bonferroni post hoc test comparing CG and EG during the exploration or intervention phase and T0-T1 and T1-T2 within groups, * p ≤ 0.05. Detailed statistics in Tables S3. (B) Mixed-effects model with Bonferroni post hoc test comparing determined and estimated cumulative energy intake. The estimated energy intake is formed as an assumed continuation of the average daily intake during the exploration phase. The Bonferroni post hoc test showed no significant differences on any study day. (C, D, and F) Linear regression, CG: n=30; EG: n=67 (some participants were excluded subsequently from this analysis because of several erroneous diary entries about unrealistic amounts of food, e.g., 30,000 g chicken for lunch). Detailed statistics in Tab. S6.

Revised version: Changes in reported daily caloric intake (A), cumulative caloric intake (B) and food composition (E) and their relationship to weight/BMI changes (C, D, and F). Data in (A) are normalized so that all baseline data of the exploration phase equal 1. To represent the x-fold change, data from the intervention phase are divided by data from the exploration phase of the corresponding participant. (A, E) 2-way repeated-measures ANOVA with Bonferroni post hoc test comparing CG and EG during the exploration or intervention phase and T0-T1 and T1-T2 within groups, * p ≤ 0.05. Detailed statistics in Tables S3. (B) Mixed-effects model with Bonferroni post hoc test comparing determined and estimated cumulative energy intake. The estimated energy intake is formed as an assumed continuation of the average daily intake during the exploration phase. The Bonferroni post hoc test showed no significant differences on any study day. (C, D, and F) Linear regression, CG: n=30; EG: n=67 (some participants were excluded subsequently from this analysis because of several erroneous diary entries about unrealistic amounts of food, e.g., 30,000 g chicken for lunch). Detailed statistics in Tab. S6.

Relation between weight change and other meal timing-related factors

Besides increasing regularity, restricting meals to certain times of the day can also shorten the daily window in which meals are eaten and shift meals to other times of the day. In fact, our intervention caused many EG participants to avoid occasional meals at extremely early or late times, limiting their meals to much shorter intervals of 11 h (95% CI [10.468, 11.520]) (Fig. 3A, B, Tab. S3), while those of the CG remained long with 13.7 h (95% CI [12.902, 14.562]). Additionally, compared with before and with the CG, the EG intervention resulted in the first meal of the day being delayed to a later time for many subjects (EG: 09:21 (95% CI [08:58, 09:45]) vs. CG: 08:03 (95% CI [07:32, 08:33])) and the last meal of the day being advanced to an earlier time (EG: 20:17 (95% CI [19:59, 20:35]) vs. CG: 21:31 (95% CI [21:05, 21:57])) (Fig. 3C, D, Tab. S3). However, weight changes of participants do not correlate with these changes in eating characteristics (Fig. 3B-D, Tab. S6).

Added in the revised version: Influence of other changes in eating times and of personalization of meal schedules on weight loss.

Earlier version: Influence of other changes in eating times and of personalization of meal schedules on weight loss. (A) 95th percentile of the time period during which food was ingested by 15 example participants during the exploration phase (yellow) and the intervention phase (purple). (B) Change in length of this period and its relation to weight/BMI changes. (C, D) Change in time of first (C) and last (D) caloric intake and their association with weight/BMI changes. (B-D) 2-way repeated-measures ANOVA with Bonferroni post hoc test comparing CG and EG during the exploration or intervention phase and T0-T1 and T1-T2 within groups, *** p ≤ 0.001, **** p ≤ 0.0001. Linear regression. EG: n=67, varying n-numbers result from incomplete entries in the diaries. Detailed statistics in Tables S3 and S6. (E) The relation between mid-Sleep phase and mid-Eat phase (brown) on the weekends of the exploration phase. The goodness of fit of this association improves during through EG intervention and affects both workdays and weekend days (purple). Linear regression, n=65, **** p ≤ 0.0001. Detailed statistics in Table S6. (F) Relation of the mid-Sleep phase (mSP) and mid-Eat phase (mEP) during exploration and intervention phase. The circles represent 24 h, the gray area represents the estimated sleep time, and the black arrow represents the mid-sleep phase, which is normalized for all subjects. The colored triangles represent the mid-Eat phase of each subject in relation to their mid-Sleep phase. The red arrow shows the average mid-Eat phase across all subjects. n=65. Detailed statistics in Table S9. (G) The smaller the residuals, i.e., the vertical distances of observation points from the regression line, became, the more BMI the EG subjects lost. Linear regression, n=65, ** p ≤ 0.01. Detailed statistics in Table S6. (H) Representative example of a meal diary showing a shift from meals on weekends. The purple arrows represent the calculated times, based on the events of the more frequent workdays, for the meal schedule of the intervention phase. (I) Shifts from weekend meals caused by the EG intervention. Paired t-test, n=65, *** p ≤ 0.001, **** p ≤ 0.0001. Data are normalized so that all baseline data at T0-T1 equal 0 to represent the change in hours. Detailed statistics in Table S2. (J) Large phase shifts of breakfast and dinner, equivalent to a correction of the social jetlag of meals on weekends, are related to greater weight loss. Linear regression, breakfast: n=55, lunch: n=64, dinner: n=63, ** p ≤ 0.01. Missing n-numbers are due to incomplete diary entries on weekends during the exploratory phase or because participants usually did not have a particular meal. Detailed statistics in Table S6. (K) Results of a multiple regression to investigate the influence of individual variables while holding the others constant on the BMI of EG participants. The Quantile-Quantile (QQ) plot shows the data are normally distributed (left). The multiple regression plot and R2 value shows that the multiple regression model explains the data well (middle). Additionally, the assignment of each regression beta coefficient to the independent variables is shown along with the p-values showing the influence on BMI. The data used is the same as those used in the individual regressions shown previously. Detailed statistics in Tab. S8. (L) Covariance matrix with normalized covariances for each pair of parameters. Values close to 1 (blue) or -1 (red) indicate high intertwining of the data, values close to 0 (white) indicate independence or orthogonality of the data. The assignment of the coefficients to the independent variables is the same as in (K).

Revised version: (A) 95th percentile of the time period during which food was ingested by 15 example participants during the exploration phase (yellow) and the intervention phase (purple). (B) Change in length of this period and its relation to weight/BMI changes. (C, D) Change in time of first (C) and last (D) caloric intake and their association with weight/BMI changes. (B-D) 2-way repeated-measures ANOVA with Bonferroni post hoc test comparing CG and EG during the exploration or intervention phase and T0-T1 and T1-T2 within groups, *** p ≤ 0.001, **** p ≤ 0.0001. Linear regression. EG: n=67, varying n-numbers result from incomplete entries in the diaries. Detailed statistics in Tables S3 and S6. (E) The relation between mid-Sleep phase and mid-Eat phase (brown) on the weekends of the exploration phase. The goodness of fit of this association improves during through EG intervention and affects both workdays and weekend days (purple). Linear regression, n=65, **** p ≤ 0.0001. Detailed statistics in Table S6. (F) Relation of the mid-Sleep phase (mSP) and mid-Eat phase (mEP) during exploration and intervention phase. The circles represent 24 h, the gray area represents the estimated sleep time, and the black arrow represents the mid-sleep phase, which is normalized for all subjects. The colored triangles represent the mid-Eat phase of each subject in relation to their mid-Sleep phase. The red arrow shows the average mid-Eat phase across all subjects. n=65. Detailed statistics in Table S9. (G) The smaller the residuals, i.e., the vertical distances of observation points from the regression line, became, the more BMI the EG subjects lost. Linear regression, n=65, ** p ≤ 0.01. Detailed statistics in Table S6. (H) Representative example of a meal diary showing a shift from meals on weekends. The purple arrows represent the calculated times, based on the events of the more frequent workdays, for the meal schedule of the intervention phase. (I) Shifts from weekend meals caused by the EG intervention. Paired t-test, n=65, *** p ≤ 0.001, **** p ≤ 0.0001. Data are normalized so that all baseline data at T0-T1 equal 0 to represent the change in hours. Detailed statistics in Table S2. (J) Large phase shifts of breakfast and dinner, equivalent to a correction of the social jetlag of meals on weekends, are related to greater weight loss. Linear regression, breakfast: n=55, lunch: n=64, dinner: n=63, ** p ≤ 0.01. Missing n-numbers are due to incomplete diary entries on weekends during the exploratory phase or because participants usually did not have a particular meal. Detailed statistics in Table S6. (K) Results of a multiple regression to investigate the influence of individual variables while holding the others constant on the BMI of EG participants. The Quantile-Quantile (QQ) plot shows the data are normally distributed (left). The multiple regression plot and R2 value shows that the multiple regression model explains the data well (middle). Additionally, the assignment of each regression beta coefficient to the independent variables is shown along with the p-values showing the influence on BMI. The data used is the same as those used in the individual regressions shown previously. Detailed statistics in Tab. S8. (L) Covariance matrix with normalized covariances for each pair of parameters. Values close to 1 (blue) or −1 (red) indicate high intertwining of the data, values close to 0 (white) indicate independence or orthogonality of the data. The assignment of the coefficients to the independent variables is the same as in (K).

Effect of meal schedule personalization on weight loss

A key part of our hypothesis involves the personalization of meal schedules. Consistent with the concept of social jet lag (the discrepancy between socially imposed and biological times), our data show that during the exploration phase, eating times on workdays differ from those on work-free days. Because social constraints usually matter less on non-working days, we expect eating times on these days to correspond more closely to individual circadian characteristics, such as the chronotype. The MCTQ questionnaire measures the chronotype based on the mid-sleep phase (mSP), i.e., the midpoint between sleep onset and waking up [25]. Corresponding to the mSP, we determined the mid-eat phase (mEP), which we define as the midpoint between the first and the last meal of the day. Indeed, our data show that the later the mSP of participants, the later their mEP on weekends (Fig. 3E, Tab. S6), with the mSP and mEP being almost antiphasic on average, i.e. 12 hours apart (Fig. 3F, Tab. S9). The EG intervention maintains most of this relationship. However, the residuals of the regression become smaller, resulting in a model fitting the data better (Fig. 3E, Tab. S6). Interestingly, the closer the intervention brings an observed value to the predicted value, i.e., the smaller the residual becomes, the more weight the participant loses (Fig. 3G, Tab. S6). This indicates that the personalized adjustment of eating times to the individual chronotype, i.e., the establishment of a highly accurate phase relation between mEP and mSP, leads to a particularly effective reduction of excess weight.

Another part of the hypothesis is that the adaptation of eating times to the times of hunger presumably determined by the circadian system contributes to weight loss. The calculation of individual eating times was based on the frequency of the hours at which meals were eaten on all days of the exploration phase. These days typically included more workdays than non-workdays (Fig. 3H). Consequently, for some subjects, the EG intervention created a discrepancy between what may have been physiologically optimal eating times, presumably reflected better by data from non-workdays, and the calculated eating times. As the EG participants were instructed to adhere to the calculated eating times even on non-working days, their usual eating times at the weekend were sometimes greatly shifted (Fig. 3I, Tab. S2). However, the magnitude of these shifts does not correlate with weight changes (Fig. S2O, Tab. S6). On the contrary, if the calculated eating times resulted in large shifts of the weekend meals, this also means that previously particularly large “meal jet lags” between working and non-working days were eliminated. Indeed, the more the calculated eating times led to a shift of late weekend breakfasts to earlier times and early weekend dinners to later times, the more the participants lost weight (Fig. 3J, Tab. S6).

Earlier version: The independent variables used above for simple regressions may influence each other and therefore their influence on BMI may be confounded. Therefore, a multiple regression with the most important and least multicollinear variables (ΔMTVS, ΔkJ, Δcalorie intake interval, ΔmSp-mEP relation, Δweekend breakfast, Δweekend dinner) was done to study their individual influences on the BMI of EG participants while keeping the other variables constant. According to this analysis, only the change in MTVS and also the resulting elimination of "dinner jet lags" have significant influence on the BMI (Fig. 3K, Tab. S8). The elimination of "breakfast jet lags" and the improvement of the mSP-mEP ratio did not reach significance but showed a strong tendency to influence BMI. In contrast, changes in caloric intake or the time interval during which food is consumed had no effect on the BMI. A concurrent analysis of covariance shows that MTVS is most strongly intertwined with BMI, while other variables play a role to a much lesser extent (Fig. 3L).

Revised version: The independent variables used above for simple regressions may influence each other and therefore their influence on BMI may be confounded. Therefore, a multiple regression with the most important and least multicollinear variables (DMTVS, DkJ, Dcalorie intake interval, DmSp-mEP relation, Dweekend breakfast, Dweekend dinner) was done to study their individual influences on the BMI of EG participants while keeping the other variables constant. According to this analysis, only the change in MTVS and also the resulting elimination of “dinner jet lags” have significant influence on the BMI (Fig. 3K, Tab. S8). The elimination of “breakfast jet lags” and the improvement of the mSP-mEP ratio did not reach significance but showed a strong tendency to influence BMI. In contrast, changes in caloric intake or the time interval during which food is consumed had no effect on the BMI. A concurrent analysis of covariance shows that MTVS is most strongly intertwined with BMI, while other variables play a role to a much lesser extent (Fig. 3L).

General well-being outcomes

Since stable circadian rhythms are generally associated with many aspects of health, regular entrainment of a variety of peripheral and brain clocks by the Zeitgeber food may enhance overall well-being. To test this, participants were asked to complete clinical questionnaires measuring subjective health status, sleep quality, depressed mood, and self-efficacy at T0 and at T2. In addition, chronotype was also determined. Many aspects of subjective physical well-being (SF-36 - physical health items) improved significantly during the intervention in the EG (Fig. 4A, Tab. S3), reversing previous differences between CG and EG. In contrast, there was only a slight improvement in the general health item within the CG. In addition, sleep quality (PSQI) of the EG group improved significantly during the intervention, but not of the CG group, again eliminating prior differences between CG and EG. Furthermore, mental health, including vitality, depressed mood, and self-efficacy (SF-36 mental health items, IDS-SR, SWE) also improved significantly in EG participants, but not in the CG (Fig. 4A, Tab. S3). Consistent with our aim to optimize rather than to change individual rhythms, chronotype did not change during the study (Fig. 4B, Tab. S3).

Added in the revised version: Improvement in overall well-being over the course of the study.

Earlier version: Improvement in overall well-being over the course of the study. (A) Significant improvements are observed in various aspects of physical and mental health (SF-36), sleep quality (PSQI), depressed mood (IDS-SR), and self-efficacy (SWE) in participants of the EG (colored columns), but hardly among CG participants (gray columns). Scales are either upward or downward, depending on whether low or high scores on the questionnaires mean subjectively good or poor well-being. (B) The chronotype of EG and CG participants did not change during the study. 2-way repeated-measures ANOVA with Bonferroni post hoc test comparing CG and EG during the exploration or intervention phase and T0-T1 and T1-T2 within groups, * p ≤ 0.05, ** p ≤ 0.01, *** p ≤ 0.001, **** p ≤ 0.0001. CG: n=33; EG: n=67. Detailed statistics in Tab. S3.

Revised version: (A) Significant improvements are observed in various aspects of physical and mental health (SF-36), sleep quality (PSQI), depressed mood (IDS-SR), and self-efficacy (SWE) in participants of the EG (colored columns), but hardly among CG participants (gray columns). Scales are either upward or downward, depending on whether low or high scores on the questionnaires mean subjectively good or poor well-being. (B) The chronotype of EG and CG participants did not change during the study. 2-way repeated-measures ANOVA with Bonferroni post hoc test comparing CG and EG during the exploration or intervention phase and T0-T1 and T1-T2 within groups, * p ≤ 0.05, ** p ≤ 0.01, *** p ≤ 0.001, **** p ≤ 0.0001. CG: n=33; EG: n=67. Detailed statistics in Tab. S3.

Discussion

Earlier version: In our pilot study, we show that increased regularity of the Zeitgeber food can lead to a significant reduction in body weight. Previously published data show that increased regularity of breakfast in particular, and not so much that of dinner, was identified as a predictor of weight loss [26]. In contrast, in our study, an increase in the regularity of lunch, and somewhat more strongly that of dinner, is related to weight loss. Although the diary entries of some participants show that they changed their energy intake and also the food composition of fat, carbohydrates, and proteins during the study, the observed weight loss is not related to these changes.

Revised version: In our pilot study, we show that increased regularity of the Zeitgeber food is associated to a significant reduction in body weight. Previously published data show that increased regularity of breakfast in particular, and not so much that of dinner, was identified as a predictor of weight loss [26]. In contrast, in our study, an increase in the regularity of lunch, and somewhat more strongly that of dinner, is related to weight loss. Although the diary entries concerning energy intake and also food composition of fat, carbohydrates, and proteins of some participants changed during the study, the observed weight loss is not related to these changes.

Related results have already been shown in mice, although here with a restriction to the entire active phase rather than specific eating times. Not surprisingly, mice gain weight and develop diabetes-like states under a high-fat diet [17]. It is surprising, however, that the negative physiological consequences of the high-fat diet can be almost completely prevented if the mice consume the equivalent amounts of high-fat food exclusively during their activity phase and do not receive any food during their sleep phase [17, 27].

Earlier version: There are also similar studies with humans in which the time window of food intake is limited to a certain length. This is, for example, 16:8 intermitted fasting or TRE/TRF, in which meals are taken only during specific time window, yet at unspecified times during that period. The effects of these protocols is mixed and often less pronounced than in our study, as less weight is lost [19–23, 28] or a longer period is required to achieve similar results [29] than in our study - with the caveat that participants with more overweight and older participants participated in these studies than in ours. Besides increased regularity, our intervention has similar effects on eating behavior, such as shortening the hours in which food is eaten. Interestingly, in our study these parameters are not related to the observed weight loss. However, other studies have also explored the effect of restricting individual meals to specific times. In mice, restricting the same amount of high-fat chow to two times per day reduces the development of obesity [30]. And in humans, restricting meals to three fixed, uniform times leads to similar weight loss like we observed in our study - but in combination with calorie reduction and standardized food [31].

Revised version: There are also similar studies with humans in which the time window of food intake is limited to a certain length. This is, for example, 16:8 intermitted fasting or TRE/TRF, in which meals are taken only during specific time window, yet at unspecified times during that period. The effects of these protocols is mixed and often less pronounced than in our study, as less weight is lost [19–23, 28] or a longer period is required to achieve similar results [29] than in our study - with the caveat that participants with more overweight and older participants participated in these studies than in ours. Besides increased regularity, our intervention has similar effects on eating behavior, such as shortening the hours in which food is eaten. Interestingly, in our study these parameters are not related to the observed weight loss.

Added in the revised version: Other studies have also explored the effect of restricting individual meals to specific times. In mice, restricting the same amount of high-fat chow to two times per day reduces the development of obesity [30]. And in humans, restricting meals to three fixed, uniform times leads to similar weight loss like we observed in our study - but in combination with imposed calorie reduction and standardized food [31].

The fact that our participants lost weight to a similar extent despite free choice of meals might be due to the individual tailoring of meal schedules. It is true that our data show that scheduling single meals to supposedly circadian optimal eating times of the participants is not essential for successful weight loss. Circadian clocks can adjust flexibly to Zeitgebers, and it is apparently more important to choose times for meals when people (have to) eat particularly frequently than the time that is probably genetically determined. That is, if the intervention leads to eating at times that are far from those normally chosen on non-working days, this has no negative impact on weight loss efficiency. Our data suggest that the regularity of eating times seems to be more important than the times themselves.

However, our pilot data still show that personalization of meal schedules is nevertheless useful. First, depending on the individual chronotype of the participants, their mealtime windows were earlier or later in the day. And those for whom the intervention optimized this relationship were particularly effective at losing excess weight. Second, weight loss is more successful if the eating time schedule is such that an individual who usually eats dinner at an early time on non-working days, but has dinner later on working days, consistently eats at the later time. This also tends to apply to breakfast if an individual eats breakfast earlier on working days than on non-working days and the meal plan is aligned with the early time of the working days.

Since food serves as a Zeitgeber for a large number of tissues, including several brain regions, it can be hypothesized that regular eating not only has a positive effect on body weight, but also on other aspects of well-being. A previous study by Panda et al. [18] has already shown that shortening the eating interval to 10-12 hours per day has a positive effect on the general well-being of the participants. Similarly, in our study, we also found significant improvements in different aspects of subjective well-being over the course of the study, specifically in physical and mental health, sleep quality, and self-efficacy. This shows that adhering to daily structured mealtimes can improve not only metabolic but also other physical and also psychological levels of health.

Earlier version: There are several explanatory approaches for why increased meal regularity can lead to weight loss despite equal caloric intake. For example, circadian clocks influence the composition of the microbiome, which regulates metabolic balance and body weight [32–37]. Additionally, circadian clocks help to anticipate mealtimes and provide the necessary digestive components in advance of food intake in advance [38–41]. Accordingly, the number of rhythmic genes in the liver of mice increases from ∼350 to ∼3,000 when the animals can follow their natural feeding rhythm and even to ∼5,000 when the feeding time is restricted to a certain period of time [42]. There could therefore be a change in metabolic programming that leads to a change in energy expenditure, for example.

Revised version: There are several explanatory approaches for why increased meal regularity can lead to weight loss despite equal caloric intake. One of these is that circadian clocks influence the composition of the microbiome, which regulates metabolic balance and body weight [32–37]. Additionally, circadian clocks help to anticipate mealtimes and provide the necessary digestive components in advance of food intake in advance [38–41]. Accordingly, the number of rhythmic genes in the liver of mice increases from ∼350 to ∼3,000 when the animals can follow their natural feeding rhythm and even to ∼5,000 when the feeding time is restricted to a certain period of time [42]. There could therefore be a change in metabolic programming that leads to a change in energy expenditure, for example.

Limitations

However, in our pilot study, we did not collect metabolic parameters that could provide mechanistic explanations for the observed weight loss. Therefore, the assumption of an optimized metabolism remains speculation and should be investigated in future follow-up studies, for example by occasional indirect calorimetry measurements that can provide information about changes in energy expenditure and the actual amount of food consumed in different phases of the study.

Furthermore, we cannot conclude whether the circadian rhythms in the digestive tract of the participants were indeed changed by the intervention. At the behavioral level, we see that the eating rhythms show significantly less variability during the intervention than before or in the CG. However, whether this has an effect on other rhythms, such as sleep-wake rhythms or physiological rhythms, remains unknown. It is known from animal studies that precise eating rhythms have a positive effect on endogenous rhythms and it is also known from human studies that regular meals help the body to adapt endogenous rhythms to meals and anticipate them [23, 38–41]. It can be assumed that this also applies to the participants in this study. Nevertheless, measurements must be integrated in future studies to confirm this assumption.

Another limitation of the pilot study is the self-assessment of key data on body weight, eating times, and food composition. Self-reported data from participants who are not subject to objective monitoring are prone to unreliability and errors. However, we have undertaken measures and analyses to reduce this uncertainty. Body weight could not be collected in our clinic under supervision because of the Covid-19 pandemic but was measured at T0, T1, and T2 in the virtual presence of study staff (by camera) in order to be able to objectively check this data for reliability. Regarding food intake, some participants underreported during the study, as indicated by unrealistically low kJ intakes. However, our analyses suggest that weight changes of individual participants are not related to their reliability of data recording as they show that the individual reporting behavior of each participant remained largely the same during the exploration and intervention phases, so underreporting is not indicative of a change in eating behavior and, conversely, the observed kJ differences are likely due to actual changes in eating behavior. However, when continuing the study, the transfer of photos of all meals consumed should be implemented to objectively assess the composition and quantity of food.

It was important for the design of this study that the participants were able to carry it out in their natural environment and not in a laboratory setting, in order to be able to convincingly demonstrate its applicability in everyday life. Instead of entries in the smartphone application, there would have been the option of questionnaires or 24-hour dietary recall outside the laboratory setting. These methods, for their part, harbor risks of misreporting due to subsequent misjudgments of food intake or possible forgetting of snacks or caloric drinks taken between meals [43, 44]. For the present pilot study, the food was to be recorded in real time in order to avoid these possible sources of error. In our view, in the natural environment, the limitation of self-reporting is unavoidable and must be accepted. The above-mentioned inclusion of indirect calorimetry in the study protocol can reveal errors in the self-report but will not be able to replace daily data collection by means of self-report.

However, the fact that the participants were only supervised online during the pilot study due to the pandemic may have had an influence. It cannot be ruled out that an occasional meeting in person would have improved compliance with data collection, as the participants might have felt more monitored by the study staff.

Although our evaluations suggest that this is the case, our data cannot clearly clarify whether the personalization of meal plans actually increases the success of weight loss. For this purpose, another control group is missing, in which the participants should have followed either a standardized meal plan or individual meal plans with specifically unsuitable eating times. This could be part of a future follow-up study, although it is difficult to predict how body weight would change in such a control group, as circadian clocks in the digestive organs can adapt to unusual eating times and the misalignment would disappear over time. However, compliance could suffer under such meal schedules, as it is more difficult for the control subjects to adhere to the eating times. This could lead to less successful weight loss.

In addition, after exclusion and drop-out of some participants, in the remaining evaluable 100 participants, there was a significant difference in baseline body weight and BMI between CG and EG. However, even if CG participants had been heavier from the start, it cannot be assumed that the applied sham treatment with an eating window of 18 hours during wake time would have had a stronger effect. Second, those subjects in the EG group whose baseline BMI was similar to those of the CG subjects lost significant weight during the intervention, while CG participants did not. Taking these two arguments into account, it cannot be assumed that a floor effect, in which the CG participants cannot lose any more weight due to a lower baseline weight, is the reason for the lack of weight loss in the CG. Lasty, studies show that even much larger baseline BMI differences do hardly affect the success of weight loss programs [45, 46], so we assume that this also applies to our study. During the course of the study, there were some drop-outs in both groups. Importantly, these did not differ in baseline weight from the completers, which is why it can be ruled out that the intervention was not effective in a particular BMI group. Rather, the subjects who dropped out were those who either could not continue due to illness or did not adhere to the meal plan. However, the inclusion of drop-outs in the analysis of weight progression does not fundamentally change the results.

Conclusion

Earlier version: The present study can be regarded as a pilot study, the results of which we believe justify a continuation with adjustment of the above-mentioned limitations. In summary, our data show that the increase of regularity of meals to a period of 90 min for each meal promotes a significant reduction of body weight and a significant increase of well-being within short time. Importantly, the successes achieved do not seem to be related to a reduction in calories, a change in food composition, or other eating behaviors. Additionally, we believe that this intervention poses hardly any health risks since it is most likely solely based on the optimization of metabolic processes. On the contrary, the strong increase in general well-being during the study rather indicates an improvement in health on many different levels.

Revised version: The present study can be regarded as a pilot study, the results of which we believe justify a continuation with adjustment of the above-mentioned limitations. In summary, our data show that the increase of regularity of meals to a period of 90 min for each meal is related to a significant reduction of body weight and a significant increase of well-being within short time. It could be of importance that the successes achieved are not related to the reported changes in calorie intake and food composition or other eating behaviors. If future studies confirm that no dietary reduction or change is required, this would indicate that the intervention optimizes the utilization of the usual diet. In addition, it can be assumed that compliance with the intervention is likely to be higher among those who implement it than would be the case with many other restrictive diets that require abstaining from certain foods or quantities. Additionally, we believe that this intervention poses hardly any health risks since it is most likely solely based on the optimization of metabolic processes. On the contrary, the strong increase in general well-being during the study rather indicates an improvement in health on many different levels.

Methods

Experimental approach

The study was a randomized controlled intervention trial with repeated measures investigating the effect of regularity of mealtimes on body weight and parameters related to overall well-being. Participants were blinded to the group assignment in that both control and experimental subjects were equally informed to participate in an intervention. The study was carried out at the Clinic for Psychiatry and Psychotherapy of the Ludwig Maximilian University (LMU), Munich, Germany. The study is registered at the German Clinical Trial Register with the trial number DRKS00021419 and was approved by the Ethics Committee of the LMU.

Participants

Earlier version: The participants were recruited across Germany via flyers at universities, fitness studios, adult education centers, pharmacies, grocery stores, and via Facebook® between September 2020 and August 2021. Those interested were able to contact the study staff by phone, email, or by making an entry in an online calendar tool. Participants had to be between 18-65 years old and have a BMI ≥ 22 kg/m². A requirement for participation in the study was that no other diet was concurrently followed, and that no medication was being taken regularly that could influence appetite or weight. No participants were included who knowingly suffer from a metabolic, mental, or addictive disorder. Further exclusion criteria were pregnancy, blindness, bedriddenness, and dependence on assistance with eating. Special emphasis was placed on the participants’ ability to understand the details of the study and to give written consent to participate.

Revised version: The participants were recruited across Germany via flyers at universities, fitness studios, adult education centers, pharmacies, grocery stores, and via Facebook® between September 2020 and August 2021. Those interested were able to contact the study staff by phone, email, or by making an entry in an online calendar tool. Participants had to be between 18-65 years old and have a BMI ≥ 22 kg/m2. A requirement for participation in the study was that no other diet was concurrently followed, and that no medication was being taken regularly that could influence appetite or weight. No participants were included who knowingly suffer from a metabolic, mental, or addictive disorder. Further exclusion criteria were pregnancy, blindness, bedriddenness, and dependence on assistance with eating. Special emphasis was placed on the participants’ ability to understand the details of the study and to give written consent to participate.

Study Protocol

Each participant completed a 12-week program, which consisted of the following steps (Fig. S1A):

Introductory Session and Initial Questionnaire Assessment

At the introductory session, the eligibility for possible participation in the study was determined and demographic data was collected. Then, suitable participants received a general introduction to the study, which included the theoretical background of the study and a detailed explanation of the study procedure. Subsequently, the participants were informed about their rights, the voluntary nature of their participation in the study, and data protection regulations. All participants had to sign an informed consent form for study participation and for the recording and use of their data. Participants weighed and measured their height themselves while study staff connected with them online via camera. Afterwards, the participant received questionnaires to fill out (Tab. S10). This initial data assessment is referred to as T0 in the analysis.

Two-Week Exploration Phase

The introduction was followed by a two-week exploration phase in which the participants were asked to follow their usual eating habits and to document all calorie intake occasions (main meals, snacks, caloric drinks) using the smartphone application “FDDB” as accurately as possible. Based on this data, the personalized nutritional schedule was then developed. Data from the FDDB smartphone application also allowed us to calculate individual caloric intake along with the ratio of consumed macronutrients for each day during the study. Additionally, the participants weight was measured again as described above. The end of the exploration phase is called T1 in the analysis.

Six-Week Intervention Phase

Added in the revised version: Experimental Group (EG)

Earlier version: Experimental Group (EG): Following the exploration phase, the participants received their personalized mealtime schedule, which was prepared beforehand by the study staff based on the meal diary of the exploration phase. Together with the participants, the nutritional behavior of the last two weeks was reviewed and analyzed. Special attention was paid to irregularities in mealtimes and to times of day when food was eaten particularly frequently. In this way, participants were made aware that they eat irregularly, but that there are nevertheless times when their internal clock frequently triggers hunger. This was followed by a six-week intervention phase, in which the participants were instructed to have their meals only at times according to the mealtime schedule, for each meal at least within a time window of ±30 min from the specified time. Participants were explicitly told that they were not restricted in their choice and quantity of food. During the intervention phase, participants should continue to document all calorie intake occasions using the FDDB App. In addition, participants were asked to weigh themselves weekly on the same scales they used at the start of the study and report their weight to the study staff to measure their progress. Halfway through the intervention phase, participants were contacted at least once by telephone to clarify possible questions and maintain compliance. Further, in case participants stopped making entries in their FDDB diary for two or more days, they were contacted to ask for reasons and to reinstate compliance again.

Revised version: Following the exploration phase, the participants received their personalized mealtime schedule, which was prepared beforehand by the study staff based on the meal diary of the exploration phase. Together with the participants, the nutritional behavior of the last two weeks was reviewed and analyzed. Special attention was paid to irregularities in mealtimes and to times of day when food was eaten particularly frequently. In this way, participants were made aware that they eat irregularly, but that there are nevertheless times when their internal clock frequently triggers hunger. This was followed by a six-week intervention phase, in which the participants were instructed to have their meals only at times according to the mealtime schedule, for each meal at least within a time window of ±30 min from the specified time. Participants were explicitly told that they were not restricted in their choice and quantity of food. During the intervention phase, participants should continue to document all calorie intake occasions using the FDDB App. In addition, participants were asked to weigh themselves weekly on the same scales they used at the start of the study and report their weight to the study staff to measure their progress. Halfway through the intervention phase, participants were contacted at least once by telephone to clarify possible questions and maintain compliance. Further, in case participants stopped making entries in their FDDB diary for two or more days, they were contacted to ask for reasons and to reinstate compliance again.

Added in the revised version: Control Group (CG)

Earlier version: Control Group (CG): In principle, participants in the CG followed the same procedure. However, they were given a sham treatment, for which they were asked to restrict their meals to a freely chosen 18-hour time window per day, during which they could eat what they wanted and, above all, when they wanted. It can be assumed that restricting the eating time window to 18 hours should have no effect on weight, since such a long eating time window is hardly ever exceeded by most people anyway [18] and therefore hardly any change in eating behavior is to be expected.

Revised version: In principle, participants in the CG followed the same procedure. However, they were given a sham treatment, for which they were asked to restrict their meals to a freely chosen 18-hour time window per day, during which they could eat what they wanted and, above all, when they wanted. It can be assumed that restricting the eating time window to 18 hours should have no effect on weight, since such a long eating time window is hardly ever exceeded by most people anyway [18] and therefore hardly any change in eating behavior is to be expected.

Final Questionnaire Assessment

At the end of the intervention phase, the participants were asked to fill out the same questionnaires as at the introduction session (Tab. S10). Additionally, weight was measured again while study staff connected with them online via camera. This data assessment is called T2 in the analysis.

Follow-up Assessment

Four weeks after completion of the study, the study staff contacted the participants again by telephone to inquire about their current weight and to obtain information about whether the participants had voluntarily continued the program after the end of the intervention phase. If the participants indicated that they did not continue the program, possible reasons were evaluated, and suggestions were received on how the current protocol could be improved to make it more attractive for permanent implementation in everyday life. Participants of the control group were informed at this point about their allocation and the actual question of the study. The final data assessment is called T3.

Restrictions due to the Covid-19 pandemic

Due to restrictions imposed by Covid-19, the study was carried out online. Data collection and communication was exclusively done via email, Zoom Meetings (Zoom Video Communications, Inc.), Facebook (Meta Platforms), or telephone.

Assessment of General Well-Being

To measure whether restricting meals to times when we believe the circadian system triggers hunger leads to an increase in general well-being, questionnaires were collected to measure sleep quality, physical well-being, self-efficacy, and depressive mood. Additionally, the chronotype was assessed. The questionnaires used are described in Tab. S10.

Mealtime Diary with FDDB Application

The nutrition diary of the smartphone application FDDB was used to document all calorie intake occasions. FDDB is freely available in the usual application stores and is operated by the independent food industry company FDDB Internetportale GmbH, Berlin. Users can record each calorie intake occasion with detailed information about the amount, type, and method of preparation of each food, which can be selected from a large database of different foods, which also contains data about the energy value of each product. The app automatically records the date and time of each calorie intake occasion. In case a certain food is not listed in the database, placeholder events with estimated energy values can be recorded. In addition to the participants, the study staff also had access to each account to check compliance and regularly download diary data.

Creation of Personal Meal Schedules

At the end of the exploration phase, FDDB data were downloaded, and clusters of times of meals were identified. First, for each participant the optimal number of mealtime clusters was determined by the heuristic technique elbow method. With the number of meals calculated in this way, the optimal times for each of these meals were then calculated using the k-means algorithm with the scikit-learn machine learning package for Python. The variable minutes was created to adjust the format of the clock times in the data. The k-means algorithm finds a clustering structure, that minimizes the sum squared error which measures the distance of each datapoint to its representative value. The center of each cluster was then calculated as the mean of all observations in the cluster, which was extracted in hours and minutes. Calculated times were rounded to quarter, half, three-quarter, or whole hours and visually displayed in scatterplots for discussion with participants of the EG. If the calculated times were incompatible with, for example, participants’ work schedules, they had the opportunity to adjust the times of meals accordingly. Then, an agreement was reached to eat only at these times for the next six weeks. For each meal, a time window of plus/minus 30 minutes around the specified time was granted. The participants were instructed not to advance missed meals or to make up for them at times that did not correspond to the established times, but to wait for the next established mealtime if possible. A scatterplot was also created for participants of the CG, but they were not provided with meal timings; instead, an 18-hour time window for meals was established with them.

Statistical Analyses

Statistical analyses were performed with SPSS 24, Python, GraphPad Prism 9.2.0, and R. Details about statistical tests used for specific experiments are indicated in the main text and the corresponding figure legends.

Mealtime Variability Score (MTVS)

A Python script was written to generate the MTVS, which is calculated from the deviation of each calorie intake occasion from the specified mealtime closest to the event. Similar to a previous study, each meal was given a score value based on the time deviation of the specified meal in minutes: 1 = +/- 0-15 min, 2 = +/- 16-30 min, 3 = +/- 31-45 min, 4 = +/- 46-60 min, 5 = +/- 61-75 min, 6 = +/- 76-90 min, 7 = +/- 91-105 min, 8 = +/- 106-120 min, 9 = +/- 2-3 h, 10 = +/- 3-4 h, 11 = over +/- 4 h [47]. Since the participants were given a 60-minute (±30 min) time slot for each meal, a score of up to 2 is considered very regular. Based on individual score values of each meal, an average score value can be calculated for each individual day, for each individual week or for the entire period, i.e., a daily, a weekly, or a total score. Similarly, daily, weekly, and total scores can be calculated for each type of meal, i.e., breakfast, lunch, dinner.

Sample Size and randomization

For the sample size calculation, we assumed an average of 2.5 kg of body weight loss, based on TRE studies [18, 19, 23, 28, 29], as we are not aware of any other study in which eating times were personalized and scheduled. The reported standard deviation for body weight loss in such studies is 3.2 [19]. Randomization was to be 2:1 (EG:CG) because no significant change in body weight was expected from the control treatment. The randomized assignment was carried out using a continuous list prepared by the study supervisor with the Excel® randomization function. Participant enrollment and assignment was carried out by study staff. On this basis a large effect size d = 0.8 was assumed for the calculation of the power analysis in G*Power and transformed into the effect size f = 0.4. The sample size calculation was performed a priori for two groups (EG and CG) and the main three measuring time points (T0, T1, T2) resulting in a sample size of N = 100 participants.

Within- and Between Group Comparisons across T0-T1 and T1-T2

Two-way repeated-measures ANOVA was used to compare group and time effects in the CG and EG and in the exploration (T0-T1) and intervention (T1-T2) phases, as well as effects of the interaction of group and time. Additionally, Bonferroni’s multiple comparison tests were used to compare effects of time within groups (T0-T1 vs. T1-T2 within CG or EG) and effects between groups at either the exploration or intervention phase (CG vs. EG at T0-T1 or T1-T2). This form of analysis was applied to data in Figures 1C & D, 2A & E, 3B-D, 4A & B, and S2K. If missing data points prevented a repeated-measure analysis, a mixed-effects model with Bonferroni’s multiple comparison test was applied instead. This applies to the data in Figures 1F, S2C, H & M. When comparisons were made only between exploration and intervention phases within a group, a paired t-test was applied, which is the case for data shown in Figures 3I and S2G. When comparisons were made with groups with very different n-numbers, Welch’s test was used instead of the t-test. This applied to data shown in Figure S2A.

Body Weight Development over Time

To analyze the change in body weight within the CG and EG groups over the T0-T2 period, one-way repeated measures ANOVA with Bonferroni’s multiple comparison test was applied. The division of CG and EG into continuing and discontinuing did not allow inclusion of T3 data in the repeated measures analysis. Therefore, the comparison between T2 and T3 was completed with a paired t-test. This type of analysis was applied in data of Figures 1E and S2A.

Relationships between parameters

Simple linear regression was used to determine whether two different parameters were related. This analysis was performed with data from Figures 1G & H, 2C, D & F, 3B-E, G & J, and S2G,, I, L, N & O. Relationships of more than two parameters were analyzed with multiple regression, which applies to data from Figures 1H and 3K.

Circadian Phase Distribution

The mSP was evaluated according to the MCTQ [25]. The mEP was calculated in the same way, but instead of using the time of falling asleep and the time of waking up as in the calculation of the mSP, the first and the last meal of the day were used for the calculation. The distribution of each calculated phase was analyzed using Rayleigh’s uniformity test. This analysis concerns the data shown in Figure 3F.